# Module 1

#### In this assignment, you will work with ufo sightings data.
- The data includes various data points about individual ufo sightings
- Data File(s): ufo-sightings.csv

<!-- grade_id: init, points: 0, grade: true, solution: false -->
```python
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################

import imp, os, sys
sol = imp.load_compiled("solutions", "./solutions.py")
sol.get_solutions("ufo-sightings.csv")
from nose.tools import assert_equal
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
'''
1. Import the csv module. Load and read the UFO sightings data set, from the ufo-sightings.csv file, 
into a DictReader inside a with statement.  Assume the data file is in the same directory as the code. 

Print the field names of the data set. Iterate over the reader to put the data into a list name "ufosightings".

'''

# your code here
import csv

filepath = "ufo-sightings.csv"
ufosightings = [] 

with open (filepath,'r') as csvfile:
    reader = csv.DictReader(csvfile)
    
    filednames = reader.fieldnames
    
    [ufosightings.append(row) for row in reader]
    


```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(ufosightings, sol.ufosightings)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
'''
2. How many sightings were there in total?  Put the count in "ufosightings_count" and print the result.
'''
# your code here
ufosightings_count = len(ufosightings)
print('The total number of sightings: ',ufosightings_count)


```

Output:
```
The total number of sightings:  80332
```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(ufosightings_count, sol.ufosightings_count)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
'''
3. How many sightings were there in the US?  Put the US sightings in "sightings_us" and print.

Hint: Check for ufo sightings where the country is 'us'.

'''

# your code here

sightings_us = [row for row in ufosightings if row["country"] == "us"]
##print('The total number of sightings in US ar: ',sightings_us)

```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(sightings_us,  sol.sightings_us)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
'''
4. Let's find the "fireball" sighting(s) that lasted more than ten seconds in US. 
Print the the datetime and state of each.  Put the data in "fball" and print the result.

Note: Consider only the US sightings stored in "sightings_us".

- Cast the duration in seconds to a float (decimal). 
- Check if duration is greater than 10. 
- Check if the shape is "fireball".

'''
#First, define a Python function that checks if a given duration (seconds) is "valid"
def is_valid_duration(duration_as_string):
# your code here
    try:
        duration_seconds = float(duration_as_string)  # try casting given string to flaod
        
    except ValueError:
        
        return False #return False if a ValueError is generated
    
    else:
        
        return duration_seconds > 0

fball =[row for row in sightings_us if is_valid_duration(row['duration (seconds)']) 
        and row['shape'] == 'fireball' 
        and float(row['duration (seconds)']) > 10]


```

<!-- grade_id: q4_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(fball, sol.fball)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
'''
5. Sort the above list by duration. What was the datetime and duration of the longest sighting?  
Put the sorted list in "fballsorted" and print the result.

- Cast the duration in seconds to a float (decimal). 
- Sort in reverse order.

'''

# your code here

fballsorted = sorted(fball, key=lambda x: float(x['duration (seconds)']), reverse=True)



```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(fballsorted, sol.fballsorted)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
'''
6. What state had the longest lasting "fireball"?   Put the state in "state" and print the result.

- Check if the shape is "fireball".
- Cast the duration in seconds to a float (decimal).
- Get the record with the largest (max) duration in seconds.
- Get the state for the record.

'''


state = fballsorted[0]['state']
state
```

Output:
```
'va'
```

<!-- grade_id: q6_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(state, sol.state)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q7, points: 0, grade: false, solution: true -->
```python
'''
7. Let's assume that any sighting (of any shape) of 0 seconds is insignificant. 
Write code to filter out these extraneous records and get the shortest sighting overall now.  
Put the minimum duration in "min_duration" and print the result.  
Use ufosightings
Note: Consider all sightings stored in "ufosightings".

'''

# your code here


min_duration = float(
                    min([row for row in ufosightings if is_valid_duration(row['duration (seconds)']) 
                    and float(row['duration (seconds)']) > 0], 
                    key=lambda x: float(x['duration (seconds)']))['duration (seconds)'])


```

<!-- grade_id: q7_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(min_duration, sol.min_duration)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q8, points: 0, grade: false, solution: true -->
```python
'''
8. What are the top 3 shapes sighted, and how many sightings were there for each? 

Note: Consider all sightings stored in "ufosightings".

- Create a new list "sightings_shapes" containing values from the "shape" column in ufosightings.  
- Create a new dictionary "count" with values of that column as keys and the counts as values.
- Get a list of the dictionary keys and values using the items() method.  This will return a list of key:value pairs.
Sort the list of key:value pairs in reverse order, from greatest (most sightings) to least.

Get the top 3 and store in "top3shapes".  Print the result.

'''

#Create a new list containing values from the "shape" column in ufosightings.
# your code here

sightings_shapes = [row['shape'] for row in ufosightings]

sightings_shapes = list(sightings_shapes)
#Create a new dictionary with values of that column as keys and the counts as values.
# your code here
count = {}

for value in sightings_shapes:
    if value in count:
        count[value] += 1
    else:
        count[value] = 1
        

#Get a list of the dictionary keys and values (use the items() method) and sort them in reverse order, from greatest (most sightings) to least.
counts = sorted(count.items(), key=lambda item: item[1],reverse=True)

count = {key: value for key, value in counts}

#Get and print the top 3.

top3shapes = counts[:3]

# your code here

```

<!-- grade_id: q8_test_a, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(sightings_shapes, sol.sightings_shapes)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q8_test_b, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(count, sol.count)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q8_test_c, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(top3shapes, sol.top3shapes)
print("Success!")
```

Output:
```
Success!
```

```python

```

```python

```